BloomTech Data Science

*Unit 2, Sprint 1*

---

👇 **Do not change the code in this cell.** If you're working in Google Colab, you can run this cell to install `category_encoders`.

```python
%%capture
# If you're in Colab...
import sys
if 'google.colab' in sys.modules:
    !pip install category_encoders==2.*
```

 
# Sprint Challenge: Predicting Home Prices in Iowa 🏠

The goal of this assignment is to build a regression model that can predict the sale price of a house in Aimes, Iowa.

# Directions 

This notebook contains 12 tasks, which cover the material we've learned in this sprint. Here's a summary:

- **Task 1:** Import training data.
- **Task 2 (`stretch goal`):** Create scatter plot.
- **Task 3:** Split data into feature matrix and target vector.
- **Task 4:** Split data into training and validation sets.
- **Task 5:** Establish baseline mean absolute error.
- **Task 6:** Build and train linear regression model.
- **Task 7:** Build and train ridge regression model.
- **Task 8:** Calculate mean absolute error for training and validation sets.
- **Task 9:** Calculate $R^2$ for validation set.
- **Task 10:** Make predictions based on test set.
- **Task 11 (`stretch goal`):** Get below 18,000 for test set MAE. 
- **Task 12 (`stretch goal`):** Plot model coefficients.

For each task you should do the following:

- Read the task instructions.
- Write your code in the cell below the task. Delete the `raise NotImplementedError` before your start.
- Run the testing cell below the task. If you get an error, read the error message and re-evaluate your code.

**You should limit your code to the following libraries:**

- `category_encoders`
- `numpy`
- `matplotlib`
- `pandas`
- `sklearn` 

# Tasks

## I. Wrangle Data

**Task 1:** Since real estate prices fluctuate over time, it's important to index the houses in this dataset using the year in which they were sold. First, load the dataset into a DataFrame named `df` and find the column that states the year sold. Next, modify the code below so that `Yr_Sold` column is parsed as a datetime object and set as the index.

```python
import numpy as np
import matplotlib.pyplot as plt
from category_encoders import OneHotEncoder
from sklearn.impute import SimpleImputer
from sklearn.linear_model import LogisticRegression
from sklearn.model_selection import train_test_split
from sklearn.pipeline import make_pipeline
from sklearn.preprocessing import StandardScaler
```

<!-- grade_id: cell-f40876cb08c183d1, points: 0, grade: false, solution: true -->
```python
'''T1: Import data into DataFrame `df`.'''
import pandas as pd

url_train = 'https://drive.google.com/uc?export=download&id=1tmYLQ1RwIgjI_d66PWnlXU-5D_VxtiQc'
df = pd.read_csv(url_train, parse_dates=['Yr_Sold'],
                 index_col='Yr_Sold')
# YOUR CODE HERE

```

```python
df.head()
```

Output:
```
1st_Flr_SF  Bedroom_AbvGr Bldg_Type  Bsmt_Full_Bath  \
Yr_Sold                                                           
2009-01-01         882              2      1Fam             0.0   
2009-01-01        1434              4      1Fam             1.0   
2009-01-01        2000              4      1Fam             1.0   
2009-01-01         945              3      1Fam             0.0   
2009-01-01        1008              4      1Fam             0.0   

            Bsmt_Half_Bath Central_Air Condition_1 Condition_2 Electrical  \
Yr_Sold                                                                     
2009-01-01             0.0           Y        Norm        Norm      SBrkr   
2009-01-01             0.0           Y        Norm        Norm      SBrkr   
2009-01-01             0.0           Y        Norm        Norm      SBrkr   
2009-01-01             0.0           Y        Norm        Norm      SBrkr   
2009-01-01             0.0           Y        Norm        Norm      SBrkr   

  
```

```python
df.info()
```

Output:
```
<class 'pandas.core.frame.DataFrame'>
DatetimeIndex: 2564 entries, 2009-01-01 to 2006-01-01
Data columns (total 46 columns):
 #   Column          Non-Null Count  Dtype  
---  ------          --------------  -----  
 0   1st_Flr_SF      2564 non-null   int64  
 1   Bedroom_AbvGr   2564 non-null   int64  
 2   Bldg_Type       2564 non-null   object 
 3   Bsmt_Full_Bath  2564 non-null   float64
 4   Bsmt_Half_Bath  2564 non-null   float64
 5   Central_Air     2564 non-null   object 
 6   Condition_1     2564 non-null   object 
 7   Condition_2     2564 non-null   object 
 8   Electrical      2564 non-null   object 
 9   Exter_Cond      2564 non-null   object 
 10  Exter_Qual      2564 non-null   object 
 11  Exterior_1st    2564 non-null   object 
 12  Exterior_2nd    2564 non-null   object 
 13  Foundation      2564 non-null   object 
 14  Full_Bath       2564 non-null   int64  
 15  Functional      2564 non-null   object 
 16  Gr_Liv_Area     2564 non-null   int64  
 17  Half_Bath      
```

**Task 1 Testing**

<!-- grade_id: cell-25880b9c025295c5, points: 1, grade: true, solution: false -->
```python
'''T1 Testing.'''
assert isinstance(df, pd.DataFrame), '`df` is not a DataFrame. Did you import the data?'
assert len(df) == 2564, 'DataFrame `df` has the wrong number of rows.'
assert isinstance(df.index, pd.DatetimeIndex), '`df` does not have a datetime index.'
```

**Task 2 (`stretch goal`):** Create a scatter plot that shows `'SalePrice'` as a function of `'Gr_Liv_Area'`. Be sure to label the axes of your plot.

<!-- grade_id: cell-8ee63eed57f965ea, points: 0, grade: false, solution: true -->
```python
import matplotlib.pyplot as plt

# YOUR CODE HERE

```

## II. Split Data

**Task 3:** Since you're building a model to predict the price at which a house will be sold, your target is the `'SalePrice'` column. Split the DataFrame `df` into the target vector `y` and the feature matrix `X`. 

<!-- grade_id: cell-4207f17d2e629cbd, points: 0, grade: false, solution: true -->
```python
'''T3: Split `df` into `X` and `y`.'''

target = 'SalePrice'
# YOUR CODE HERE
X = df.drop(columns=target)
y = df[target]
```

**Task 3 Testing**

<!-- grade_id: cell-7e83bdb6a35f2699, points: 1, grade: true, solution: false -->
```python
'''T3 Tests'''
assert y.shape == (2564,), '`y` either has the wrong number of rows, or is two-dimentional.'
assert X.shape == (2564, 45), '`X` has the wrong number of columns or rows.'
assert 'SalePrice' not in X.columns, 'Target is still part of feature matrix.' 
```

**Task 4:** Split `X` and `y` into training and validation sets. Since housing prices fluctuate over time, you should use the time-based cutoff method. 

- Training: Homes sold before 2009
- Validation: Homes sold during 2009

<!-- grade_id: cell-2efa29975d55a10c, points: 0, grade: false, solution: true -->
```python
'''T4: Split `X` and `y` into training, validation, and testing sets.'''

# YOUR CODE HERE
cutoff = '2009'
mask = X.index < cutoff
X_train, y_train = X.loc[mask], y.loc[mask]
X_val, y_val = X.loc[~mask], y.loc[~mask]

```

**Task 4 Testing**

<!-- grade_id: cell-458101f90b6b100b, points: 1, grade: true, solution: false -->
```python
'''T4 Tests'''
assert len(X_train) == len(y_train) == 1920, 'Your training set has the wrong number of rows.'
assert len(X_val) == len(y_val) == 644, 'Your validation set has the wrong number of rows.'
```

## III. Set Baseline

**Task 5:** Before you can create a model to predict sale price, you need to establish a baseline. First, calculate the mean housing price for your training set. Next, calculate the *mean absolute error* for a model that always predicts that mean. Be sure to assign the mean absolute error to a variable named `baseline_mae`.

```python
y_train.head()
```

Output:
```
Yr_Sold
2008-01-01    176500
2008-01-01    237500
2008-01-01    206900
2008-01-01    187500
2008-01-01    165000
Name: SalePrice, dtype: int64
```

<!-- grade_id: cell-62dc6e7707644a97, points: 0, grade: false, solution: true -->
```python
'''T5: Calculate baseline mean absolute error.'''
from sklearn.metrics import mean_absolute_error

# YOUR CODE HERE
y_train_pred = [y_train.mean()] * len(y_train)
baseline_mae = mean_absolute_error(y_train, y_train_pred)
print('Baseline MAE:', baseline_mae)
```

Output:
```
Baseline MAE: 58502.9394639757
```

**Task 5 Testing**

<!-- grade_id: cell-cfcf0f63baa004f9, points: 1, grade: true, solution: false -->
```python
'''T5 Tests'''
assert isinstance(baseline_mae, float), '`baseline_mae` is not a `float`. Have you assigned a value to this variable name?'
```

## IV. Build Models

**Task 6:** Build and train a linear regression model named `model_lr` that predicts home sale price. Your model should have the following components:

- A `OneHotEncoder` for categorical features.
- A `StandardScaler` so that you can compare your model coefficients after training.
- A `pipeline` that combines transformers and predictor.

**Tip:** For your `OneHotEncoder`, `use_cat_names` should be set to `True`.

<!-- grade_id: cell-b80a918d2acd5e24, points: 0, grade: false, solution: true -->
```python
'''T6: Build and train LinearRegression model.'''
from sklearn.pipeline import make_pipeline
from sklearn.linear_model import LinearRegression
from category_encoders import OneHotEncoder
from sklearn.preprocessing import StandardScaler

# YOUR CODE HERE
model_lr = make_pipeline(
    OneHotEncoder(use_cat_names=True),
    SimpleImputer(),
    StandardScaler(),
    LinearRegression()
)
model_lr.fit(X_train, y_train);
```

**Task 6 Testing**

<!-- grade_id: cell-066938e1bd64a34c, points: 1, grade: true, solution: false -->
```python
'''T6 Tests'''
from sklearn.pipeline import Pipeline
assert isinstance(model_lr, Pipeline), '`model_lr` is not a pipeline.'
assert len(model_lr.named_steps) > 1, 'Your `model_lr` pipeline is missing components.'
```

```python
df.head()
```

Output:
```
1st_Flr_SF  Bedroom_AbvGr Bldg_Type  Bsmt_Full_Bath  \
Yr_Sold                                                           
2009-01-01         882              2      1Fam             0.0   
2009-01-01        1434              4      1Fam             1.0   
2009-01-01        2000              4      1Fam             1.0   
2009-01-01         945              3      1Fam             0.0   
2009-01-01        1008              4      1Fam             0.0   

            Bsmt_Half_Bath Central_Air Condition_1 Condition_2 Electrical  \
Yr_Sold                                                                     
2009-01-01             0.0           Y        Norm        Norm      SBrkr   
2009-01-01             0.0           Y        Norm        Norm      SBrkr   
2009-01-01             0.0           Y        Norm        Norm      SBrkr   
2009-01-01             0.0           Y        Norm        Norm      SBrkr   
2009-01-01             0.0           Y        Norm        Norm      SBrkr   

  
```

**Task 7:** Build and train a ridge regression model named `model_r` that predicts home sales price. Your model should include the following components:

- A `OneHotEncoder` for categorical features.
- A `StandardScaler` so that you can compare your model coefficients after training.
- A `pipeline` that combines transformers and predictor.

**Tip:** If you combine a `StandardScaler` transformer with a `Ridge` predictor, you should set the `normalize` argument for the latter to `False`.

<!-- grade_id: cell-02455b6368014e2e, points: 0, grade: false, solution: true -->
```python
'''T7. Build and train Ridge model.'''
from sklearn.linear_model import Ridge
from sklearn.pipeline import make_pipeline
from sklearn.pipeline import Pipeline
# YOUR CODE HERE
model_r = make_pipeline(
    OneHotEncoder(use_cat_names=True),
    StandardScaler(),
    Ridge(normalize=False)
)
 
model_r.fit(X_train, y_train);
```

Output:
```
/usr/local/lib/python3.8/dist-packages/sklearn/linear_model/_base.py:148: FutureWarning: 'normalize' was deprecated in version 1.0 and will be removed in 1.2. Please leave the normalize parameter to its default value to silence this warning. The default behavior of this estimator is to not do any normalization. If normalization is needed please use sklearn.preprocessing.StandardScaler instead.
  warnings.warn(
```

**Task 7 Testing**

<!-- grade_id: cell-587e7ab699c005cb, points: 1, grade: true, solution: false -->
```python
assert isinstance(model_r, Pipeline), '`model_r` is not a pipeline.'
assert len(model_r.named_steps) > 1, 'Your `model_r` pipeline is missing components.'
```

## V. Check Metrics

**Task 8:** Compare your models by calculating their mean absolute error for the training and validation sets. You should assign these to the variables `model_lr_training_mae`, `model_lr_validation_mae`, `model_r_training_mae`, and `model_r_validation_mae`, respectively. The validation MAE for your Ridge model should be lower than your `baseline_mae`.

```python
df.head()
```

Output:
```
1st_Flr_SF  Bedroom_AbvGr Bldg_Type  Bsmt_Full_Bath  \
Yr_Sold                                                           
2009-01-01         882              2      1Fam             0.0   
2009-01-01        1434              4      1Fam             1.0   
2009-01-01        2000              4      1Fam             1.0   
2009-01-01         945              3      1Fam             0.0   
2009-01-01        1008              4      1Fam             0.0   

            Bsmt_Half_Bath Central_Air Condition_1 Condition_2 Electrical  \
Yr_Sold                                                                     
2009-01-01             0.0           Y        Norm        Norm      SBrkr   
2009-01-01             0.0           Y        Norm        Norm      SBrkr   
2009-01-01             0.0           Y        Norm        Norm      SBrkr   
2009-01-01             0.0           Y        Norm        Norm      SBrkr   
2009-01-01             0.0           Y        Norm        Norm      SBrkr   

  
```

```python
list(df.columns)
```

Output:
```
['1st_Flr_SF',
 'Bedroom_AbvGr',
 'Bldg_Type',
 'Bsmt_Full_Bath',
 'Bsmt_Half_Bath',
 'Central_Air',
 'Condition_1',
 'Condition_2',
 'Electrical',
 'Exter_Cond',
 'Exter_Qual',
 'Exterior_1st',
 'Exterior_2nd',
 'Foundation',
 'Full_Bath',
 'Functional',
 'Gr_Liv_Area',
 'Half_Bath',
 'Heating',
 'Heating_QC',
 'House_Style',
 'Kitchen_AbvGr',
 'Kitchen_Qual',
 'Land_Contour',
 'Land_Slope',
 'Lot_Area',
 'Lot_Config',
 'Lot_Shape',
 'MS_SubClass',
 'MS_Zoning',
 'Mas_Vnr_Type',
 'Mo_Sold',
 'Neighborhood',
 'Overall_Cond',
 'Overall_Qual',
 'Paved_Drive',
 'Roof_Matl',
 'Roof_Style',
 'SalePrice',
 'Sale_Condition',
 'Sale_Type',
 'Street',
 'TotRms_AbvGrd',
 'Utilities',
 'Year_Built',
 'Year_Remod/Add']
```

```python
df['Condition_1'].value_counts()
```

Output:
```
Norm      2209
Feedr      142
Artery      77
RRAn        46
PosN        36
RRAe        25
PosA        17
RRNn         8
RRNe         4
Name: Condition_1, dtype: int64
```

<!-- grade_id: cell-10ea192c7e6ee9de, points: 0, grade: false, solution: true -->
```python
'''T8: Compare model metrics, MAE.'''
# YOUR CODE HERE
model_lr_training_mae = mean_absolute_error(y_train, model_lr.named_steps['linearregression'].predict(X_train))

print('Linear Regression Model')
print('Training MAE:', model_lr_training_mae)
print('Validation MAE:', model_lr_validation_mae)
print()
print('Ridge Regression Model')
print('Training MAE:', model_r_training_mae)
print('Validation MAE:', model_r_validation_mae)
```

Output:
```
/usr/local/lib/python3.8/dist-packages/sklearn/base.py:443: UserWarning: X has feature names, but LinearRegression was fitted without feature names
  warnings.warn(
/usr/local/lib/python3.8/dist-packages/sklearn/utils/validation.py in check_array(array, accept_sparse, accept_large_sparse, dtype, order, copy, force_all_finite, ensure_2d, allow_nd, ensure_min_samples, ensure_min_features, estimator)
    744                     array = array.astype(dtype, casting="unsafe", copy=False)
    745                 else:
--> 746                     array = np.asarray(array, order=order, dtype=dtype)
    747             except ComplexWarning as complex_warning:
    748                 raise ValueError(

/usr/local/lib/python3.8/dist-packages/pandas/core/generic.py in __array__(self, dtype)
   1991 
   1992     def __array__(self, dtype: NpDtype | None = None) -> np.ndarray:
-> 1993         return np.asarray(self._values, dtype=dtype)
   1994 
   1995     def __array_wrap__(

ValueError: could not
```

**Task 8 Testing**

<!-- grade_id: cell-e467cc19a51d9a41, points: 1, grade: true, solution: false -->
```python
'''T8 Tests'''
assert isinstance(model_lr_training_mae, float), '`model_lr_training_mae` is not a float. Did you assign a value to this variable?'
assert isinstance(model_lr_validation_mae, float), '`model_lr_validation_mae` is not a float. Did you assign a value to this variable?'
assert isinstance(model_r_training_mae, float), '`model_r_training_mae` is not a float. Did you assign a value to this variable?'
assert isinstance(model_r_validation_mae, float), '`model_r_validation_mae` is not a float. Did you assign a value to this variable?'
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-32-5d7760038a10> in <module>
      1 '''T8 Tests'''
----> 2 assert isinstance(model_lr_training_mae, float), '`model_lr_training_mae` is not a float. Did you assign a value to this variable?'
      3 assert isinstance(model_lr_validation_mae, float), '`model_lr_validation_mae` is not a float. Did you assign a value to this variable?'
      4 assert isinstance(model_r_training_mae, float), '`model_r_training_mae` is not a float. Did you assign a value to this variable?'
      5 assert isinstance(model_r_validation_mae, float), '`model_r_validation_mae` is not a float. Did you assign a value to this variable?'

NameError: name 'model_lr_training_mae' is not defined
```

**Task 9:** Compare your two models by calculating their $R^2$ score for the **validation data**.

<!-- grade_id: cell-92d55f4de3bb7f31, points: 0, grade: false, solution: true -->
```python
'''T9: Compare model metrics, R^2.'''
# YOUR CODE HERE
raise NotImplementedError()
print('Linear Regression R^2:', model_lr_r2_score)
print('Ridge Regression R^2:', model_r_r2_score)
```

**Task 9 Testing**

<!-- grade_id: cell-7b79f2550b912074, points: 1, grade: true, solution: false -->
```python
'''T9 Tests'''
assert isinstance(model_lr_r2_score, float), '`model_lr_r2_score` is not a float. Did you assign a value to this variable?'
assert isinstance(model_r_r2_score, float), '`model_r_r2_score` is not a float. Did you assign a value to this variable?'
assert model_lr_r2_score <= 1, 'The maximum value for an R^2 score is 1.'
assert model_r_r2_score <= 1, 'The maximum value for an R^2 score is 1.'
```

**Task 10:** Based on the metrics above, choose the best of your two models and generate an array of predictions `y_pred` using the test set below.

**Note:** Make sure that you import and clean the test set *exactly the same way* you did with the training set at the start of this notebook.

<!-- grade_id: cell-87695ab01935fdd9, points: 0, grade: false, solution: true -->
```python
url_test = 'https://drive.google.com/uc?export=download&id=1y9u8cOWprTjruw8E-ct1c7YaEytToqd_'
X_test = pd.read_csv(url_test)

# YOUR CODE HERE
raise NotImplementedError()
print('My predictions:', y_pred[:3])
```

**Task 10 Testing:** There is a hidden test in the cell below. When you submit your notebook to Canvas, your test MAE will be calculated. In order to get the point for this task, your test MAE must be below 58,000.

**Tip:** Use your validation MAE to estimate what your test MAE will be.

<!-- grade_id: cell-c3c94d263dabfe02, points: 1, grade: true, solution: false -->
```python
import numpy as np
assert isinstance(y_pred, np.ndarray), '`y_pred` should be type `ndarray`.'
assert y_pred.shape == (340,), '`y_pred` should be a 1-dimentional array with 340 predictions.'
# Your test MAE will be printed out below when you submit your assignment.
```

**Task 11 (`stretch goal`):** There's no code to write for this task. Below is another testing cell that will calculate the MAE for your test set. In order to get the point for this task, your test MAE must be below 18,000.

**Tip:** Are there any **hyperparameters** in your model that you could adjust to improve performance? 

<!-- grade_id: cell-40815d96e2a2adf9, points: 0, grade: true, solution: false -->
```python
assert isinstance(y_pred, np.ndarray), '`y_pred` should be type `ndarray`.'
assert y_pred.shape == (340,), '`y_pred` should be a 1-dimentional array with 340 predictions.'
# Your test MAE will be printed out below when you submit your assignment.
```

## VI. Explain Model

**Task 12 (`stretch goal`):** Create a horizontal barchart for your best performing model that shows the ten most important features for the model's predictions. Don't forget to label your axes!

This task will not be autograded - but it is part of completing the challenge.

<!-- grade_id: cell-ebd0b1ea2eab31c6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```